Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Pchelintsev Ilya"
COLLABORATORS = ""
```

---

# Интеграл от осциллирующей функции

Рассмотрим определенный интеграл

$$
I = \int_0^\infty\! \frac{\sin(x) \cos{(\cos{(x)})}}{x}\,dx
$$

Вычислите значение данного интеграла с относительной точностью $10^{-6}$. Для упрощения задачи можете воспользоваться функционалом `scipy.integrate.quad`.

Заметим, что "из коробки" интеграл вычислить не удается, и нужно что-то придумать.

```python
from math import sin, cos

from scipy.integrate import quad
quad(lambda x: sin(x) * cos(cos(x)) / x, 0, float('inf'))
```

Output:
```
<ipython-input-23-5a28b85f551f>:4: IntegrationWarning: The integral is probably divergent, or slowly convergent.
  quad(lambda x: sin(x) * cos(cos(x)) / x, 0, float('inf'))
(1.9653912540956733, 4.089174284042322)
```

Напишите функцию, которая возвращает значение данного интеграла. Только само значение интеграла, без оценки погрешности.
(оборачиваем интеграл в функцию только для удобства автопроверки).

<!-- grade_id: cell-ddd4e0b6d148fcfd, points: 0, grade: false, solution: true -->
```python
from scipy.integrate import quad

def integ():
    """Return a single float, the computed value of the integral."""
    eps = 1e-7
    x0 = eps
    dx = 100
    I = 0.
    I_old = -1.
    n=0
    while True:
        I_0 = quad(lambda x: sin(x) * cos(cos(x)) / x, x0, x0 + dx)
        I += I_0[0]
        x0 += dx
        if abs(I_0[0]-I_old) < eps and abs(I_0[0]) < eps:
            break
        I_old = I_0[0]
        n+=1
    return I
```

<!-- grade_id: cell-b341c2ce36bc957d, points: 5, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

# this is a test cell, keep it intact

```

<!-- grade_id: cell-b1efbd0432fcd7cb, points: 15, grade: true, solution: false -->
```python
# this is a test cell, keep it intact

```